Early menarche and childbirth accelerate aging-related outcomes and age-related diseases: Evidence for antagonistic pleiotropy in humans

Author response:

The following is the authors’ response to the previous reviews

Reviewer #1 (Public review):

Summary:

The present study aims to associate reproduction with age-related disease as support of the antagonistic pleiotropy hypothesis of ageing predominantly using Mendelian Randomization. The authors found evidence that early-life reproductive success is associated with advanced ageing.

Strengths:

Large sample size. Many analyses.

Weaknesses:

Still a number of doubts with regard to some of the results and their interpretation.

Reviewer #1 (Recommendations for the authors):

Thank you for the opportunity to review a revised version.

I still have serious doubts with regard to a number of datasets presented. For example, the results on essential hypertension and cervical cancer show very small effect sizes, but according to the authors still reach the level of statistical significance. This is unlikely to be accurate. For MR analyses, this is nearly impossible. The analyses of these data and the statistical analysis need to be checked for errors and repeated. While BOLT-LLM might not be relevant here, there might be other things happening here. The authors should therefore always interpret the results also with regard to the observed effect sizes instead of only looking at the p-values (0.999 means that there is a 0.1% lower risk).

Thank you for your suggestions. We have updated the results for essential hypertension, GAD, and cervical cancer in results, figures, and supplemental tables (lines 65-89, Figure 1, Tables S3-S4).

Reviewer #2 (Public review):

Summary:

The authors present an interesting paper where they test the antagonistic pleiotropy theory. Based on this theory they hypothesize that genetic variants associated with later onset of age at menarche and age at first birth may have a positive effect on a multitude of health outcomes later in life, such as epigenetic aging and prevalence of chronic diseases. Using a mendelian randomization and colocalization approach, the authors show that SNPs associated with later age at menarche are associated with delayed aging measurements, such as slower epigenetic aging and reduced facial aging and a lower risk of chronic diseases, such as type 2 diabetes and hypertension. Moreover, they identify 128 fertility-related SNPs that associate with age-related outcomes and they identified BMI as a mediating factor for disease risk, discussing this finding in the context of evolutionary theory.

Strengths:

The major strength of this manuscript is that it addresses the antagonistic pleiotropy theory in aging. Aging theories are not frequently empirically tested although this is highly necessary. The work is therefore relevant for the aging field as well as beyond this field, as the antagonistic pleiotropy theory addresses the link between fitness (early life health and reproduction) and aging.

The authors addressed the remarks on the previous version very well. Addressing the two points below would further increase the quality of the manuscript.

(1) In the previous version the authors mentioned that their results are also consistent with the disposable soma theory: "These results are also consistent with the disposable soma theory that suggests aging as an outcome tradeoff between an organism's investment in reproduction and somatic maintenance and repair."

Although the antagonistic pleiotropy and disposable soma theories describe different mechanisms, both provide frameworks for understanding how genes linked to fertility influence health. The antagonistic pleiotropy theory posits that genes enhancing fertility early in life may have detrimental effects later. In contrast, the disposable soma theory suggests that energy allocation involves a trade-off, where investment in fertility comes at the expense of somatic maintenance, potentially leading to poorer health in later life.

To strengthen the manuscript, a discussion section should be added to clarify the overlap and distinctions between these two evolutionary theories and suggest directions for future research in disentangling their specific mechanisms.

Thank you for your suggestions to clarify the overlap and distinctions between the antagonistic pleiotropy and disposable soma theories. While our primary focus is on the antagonistic pleiotropy framework, we acknowledge that the disposable soma theory also provides a relevant perspective on the trade-offs between reproduction and somatic maintenance.

To address this, we have expanded the discussion section to highlight how both theories contribute to our understanding of the relationship between fertility-related traits and aging-related health outcomes. We also suggested potential future research directions, such as integrating genetic data with biomarkers of somatic to further explore the mechanisms underlying these trade-offs (lines 213-223).

(2) In response to the question why the authors did not include age at menopause in addition to the already included age at first child and age at menarche the following explanation was provided: "Our manuscript focuses on the antagonistic pleiotropy theory, which posits that inherent trade-off in natural selection, where genes beneficial for early survival and reproduction (like menarche and childbirth) may have costly consequences later. So, we only included age at menarche and age at first childbirth as exposures in our research."

It remains, however, unclear why genes beneficial for early survival and reproduction would be reflected only in age at menarche and age at first childbirth, but not in age at menopause. While age at menarche marks the onset of fertility, age at menopause signifies its end. Since evolutionary selection acts directly until reproduction is no longer possible (though indirect evolutionary pressures persist beyond this point), the inclusion of additional fertility-related measures could have strengthened the analysis. A more detailed justification for focusing exclusively on age at menarche and first childbirth would enhance the clarity and rigor of the manuscript.

Thank you for your question regarding the age at menopause in our analysis. Our decision was based on the theoretical framework of antagonistic pleiotropy, which emphasizes early-life reproductive advantages that may have trade-offs later in life. Age at menarche and age at first childbirth are direct markers of early reproductive investment, which align closely with this framework.

While age at menopause marks the cessation of reproductive capability, its evolutionary role is distinct. The selective pressures acting on menopause are complex and may involve post-reproductive contributions rather than direct reproductive fitness benefits. Moreover, the genetic architecture of menopause may be influenced by different biological pathways compared to early reproductive traits.

Nonetheless, we acknowledge that including age at menopause could provide additional insights into reproductive aging. Several papers1,2 were already published regarding age at menopause and age-related outcomes, including diabetes, AD, osteoporosis, cancers, and cardiovascular diseases.

Reviewing Editor (Recommendations for the authors):

Above/below you will find the remaining comments from the reviewers. One of the main issues remaining is that some of the data seems to be incorrectly analysed and some of the findings may not be correct. To clarify this a lot more, I asked the reviewer for some details and received the following:

- In Figure 1B one of their main outcomes is "age of menopause", but they report the data as an odds ratio. This is not correct and should be fixed (it seems the authors can run the right analysis, but just reported it with the wrong heading in the figure). This likely also applies to the outcome "facial aging". Also the heading in Figure 1A should be Beta instead of OR.

We have updated the figures to ensure that the beta values of continuous outcomes and odds ratio values of categorical outcomes are presented in Figure 1.

- With essential hypertension, GAD and cervical cancer, the estimates are so small that they need to re-review their results. The current MR analysis is not sufficiently powered to have such small confidence intervals. Essential hypertension was based on data from UK biobank, although I was also unable to find what program was used to generate the GWAS results, I have strong thoughts this was also BOLT-LLM. Same for cervical cancer. Both datasets used familial-related samples, so they are very likely derived with BOLT-LLM.

I hope this will help to solve this issue.

Based on published paper, gastrointestinal or abdominal disease (GAD) (GWAS ID: ebi-a-GCST90038597) is after BOLT-LLM. Based on MRC IEU UK Biobank GWAS pipeline, version 1 and 2, essential hypertension (GWAS ID: ukb-b-12493) and cervical cancer (GWAS ID: ukb-b-8777) are after BOLT-LLM. We have updated the MR analysis results and figures (lines 65-89, Figure 1, Tables S3-S4) as well as the following IPA analysis (lines 106-162 and 255-280, Figures 2-3).

(1) Magnus, M. C., Borges, M. C., Fraser, A. & Lawlor, D. A. Identifying potential causal effects of age at menopause: a Mendelian randomization phenome-wide association study. Eur J Epidemiol 37, 971-982 (2022). https://doi.org:10.1007/s10654-022-00903-3

(2) Zhang, X., Huangfu, Z. & Wang, S. Review of mendelian randomization studies on age at natural menopause. Front Endocrinol (Lausanne) 14, 1234324 (2023). https://doi.org:10.3389/fendo.2023.1234324

Changes from the earlier version to the revised version

Abstract

Abstract

Earlier version: Aging can be understood as a consequence of the declining force of natural selection with age. Consistent with this, the antagonistic pleiotropy theory of aging proposes that aging arises from trade-offs that favor early growth and reproduction. However, evidence supporting antagonistic pleiotropy in humans remains limited. Using Mendelian Randomization (MR), we demonstrated that later ages of menarche or first childbirth were genetically associated with longer parental lifespan, decreased frailty index, slower epigenetic aging, later menopause, and reduced facial aging. Moreover, later menarche or first childbirth were also genetically associated with a lower risk of several age-related diseases, including late-onset Alzheimer’s disease (LOAD), type 2 diabetes, heart disease, essential hypertension, and chronic obstructive pulmonary disease (COPD). We validated the associations between the age of menarche, childbirth, and the number of childbirths with several age-related outcomes in the UK Biobank by conducting regression analysis of nearly 200,000 subjects. Our results demonstrated that menarche before the age 11 and childbirth before 21 significantly accelerated the risk of several diseases, and almost doubled the risk for diabetes, heart failure, and quadrupled the risk of obesity, supporting the antagonistic pleiotropy theory. We identified 126 significant single nucleotide polymorphisms (SNPs) that influenced age-related outcomes, some of which were involved in known longevity pathways, including IGF1, growth hormone, AMPK, and mTOR signaling. Our study also identified higher BMI as a mediating factor in causing the increased risk of certain diseases, such as type 2 diabetes and heart failure, in women with early menarche or early pregnancy, emphasizing the importance of the thrifty gene hypothesis in explaining in part the mechanisms behind antagonistic pleiotropy. Our study highlights the complex relationship between genetic legacies and modern diseases, emphasizing the need for gender-sensitive healthcare strategies that consider the unique connections between female reproductive health and aging.

Revised version: Aging can be understood as a consequence of the declining force of natural selection with age. Consistent with this, the antagonistic pleiotropy theory of aging proposes that aging arises from trade-offs that favor early growth and reproduction. However, evidence supporting antagonistic pleiotropy in humans remains limited. Using Mendelian Randomization (MR), we demonstrated that later ages of menarche or first childbirth were genetically associated with longer parental lifespan, decreased frailty index, slower epigenetic aging, later menopause, and reduced facial aging. Moreover, later menarche or first childbirth were also genetically associated with a lower risk of several age-related diseases, including late-onset Alzheimer’s disease (LOAD), type 2 diabetes, heart disease, essential hypertension, and chronic obstructive pulmonary disease (COPD). We validated the associations between the age of menarche, childbirth, and the number of childbirths with several age-related outcomes in the UK Biobank by conducting regression analysis of nearly 200,000 subjects. Our results demonstrated that menarche before the age 11 and childbirth before 21 significantly accelerated the risk of several diseases, and almost doubled the risk for diabetes, heart failure, and quadrupled the risk of obesity, supporting the antagonistic pleiotropy theory. We identified 158 significant single nucleotide polymorphisms (SNPs) that influenced age-related outcomes, some of which were involved in known longevity pathways, including IGF1, growth hormone, AMPK, and mTOR signaling. Our study also identified higher BMI as a mediating factor in causing the increased risk of certain diseases, such as type 2 diabetes and heart failure, in women with early menarche or early pregnancy, emphasizing the importance of the thrifty gene hypothesis in explaining in part the mechanisms behind antagonistic pleiotropy. Our study highlights the complex relationship between genetic legacies and modern diseases, emphasizing the need for gender-sensitive healthcare strategies that consider the unique connections between female reproductive health and aging.

Introduction

Given the decline in the force of natural selection with age, genes that influence aging and age-related diseases are likely to be selected for their influence on early life events 1. The theory of antagonistic pleiotropy suggests that natural selection involves inherent trade-offs, where genes that are advantageous for early survival and reproduction may have detrimental effects later in life, contributing to the aging process and the development of age-related diseases 1-3. Though evidence for antagonistic pleiotropy has been observed in invertebrate models, evidence for causal relationships in mammals, especially humans, is largely lacking 4. The timing of reproductive events, such as menarche and childbirth, has long been recognized as a crucial aspect of human life history and evolution. However, accumulating evidence suggests that these reproductive milestones may have far-reaching implications 5,6.

Considering reproductive events are partly regulated by genetic factors that can manifest the physiological outcome later in life, we hypothesized that earlier or later onset of menarche and first childbirth could reflect broader genetic influences on longevity and disease susceptibility, serving as proxies for biological aging processes (fig. S1 and fig. S2). Mendelian randomization (MR) is a term that applies to the use of genetic variation to address causal questions about how modifiable exposures influence different outcomes 7. Two-sample and two-step MR analyses were adopted in our research to explore the genetic causal associations using the inverse variance weighted (IVW) model. We leveraged recent studies that have identified SNPs associated with early menarche or birth in humans, to examine their role in antagonistic pleiotropy. After preprocessing, there were 209 and 33 single nucleotide polymorphisms (SNPs) included in MR analyses for the two exposures, age at menarche 8 and age at first birth 9, respectively.

Later menarche was genetically associated with later aging outcomes

Earlier version: Since parental ages at death may reflect the genetic heritability of lifespan, parental ages at death were used as general aging outcomes 10. Compared to early menarche, later age at menarche was significantly associated with later parental ages at death for mothers or fathers (Beta=2.27E-02, P=1.84×10−4; Beta=1.97E-02, P=7.88×10−4 for father and mother’s ages at death respectively) (Fig. 1A and Table S1-S3). Aging is the biggest risk factor for frailty and several age-related diseases. Later age at menarche was associated with a lower frailty index, which was calculated using a set of 49 self-reported questionnaire items on traits covering health, presence of diseases and disabilities, and mental well-being 11 (Beta=-2.36E-02, P=5.75×10−3). After excluding the instrumental variables (IVs) associated with body mass index (BMI), the significant association between later menarche and outcomes remained based on the weighted median model (WMM) for frailty index (Beta=-2.88E-02, P=0.0202) and based on IVW for parental ages at death (Beta=1.47E-02, P=2.46×10−2; Beta=1.24E-02, P=3.87×10−2). Next, we examined the associations between the age of menarche and organ aging or age-related diseases (Fig. 1B). Later age of menarche was associated with later menopause onset (Beta=1.16E-01, P=2.11×10−2 with WMM), lower risks of late-onset Alzheimer’s disease (LOAD) (OR=0.897 (0.810–0.994), P=3.88×10−2), chronic heart failure (CHF) (OR=0.953 (0.913–0.995), P=2. 80×10−2), essential hypertension (OR=0.990 (0.986–0.993), P=1.99×10−8), facial aging (Beta=-1.76E-02, P=2.16×10−9), and early onset chronic obstructive pulmonary disease (COPD) (OR=0.838 (0.755–0.931), P=9.50×10−4) and mildly higher risk of osteoporosis (OR=1.009 (1.006–1.013), P=1.95×10−7). These associations were still significant after excluding SNPs associated with BMI, except for CHF. One potential mechanism of antagonistic pleiotropy is accelerated cell growth and division; therefore, we examined the relationship between menarche and certain cancers. Later age at menarche was associated with lower risks of breast cancer (OR=0.859 (0.764–0.966), P=1.53×10−2) and endometrial cancer (OR=0.896 (0.840–0.956), P=9.02×10−4) compared with early menarche (Fig. 1B).

Revised version: Since parental ages at death may reflect the genetic heritability of lifespan, parental ages at death were used as general aging outcomes 10. Compared to early menarche, later age at menarche was significantly associated with later parental ages at death for mothers or fathers (Beta=2.27E-02, P=1.84×10−4; Beta=1.97E-02, P=7.88×10−4 for father and mother’s ages at death respectively) (Fig. 1A and Table S1-S3). Aging is the biggest risk factor for frailty and several age-related diseases. Later age at menarche was associated with a lower frailty index, which was calculated using a set of 49 self-reported questionnaire items on traits covering health, presence of diseases and disabilities, and mental well-being 11 (Beta=-2.36E-02, P=5.75×10−3). After excluding the instrumental variables (IVs) associated with body mass index (BMI), the significant association between later menarche and outcomes remained based on the weighted median model (WMM) for frailty index (Beta=-2.88E-02, P=0.0202) and based on IVW for parental ages at death (Beta=1.47E-02, P=2.46×10−2; Beta=1.24E-02, P=3.87×10−2). Next, we examined the associations between the age of menarche and organ aging or age-related diseases (Fig. 1A and 1B and Table S3). Later age of menarche was associated with later menopause onset (Beta=1.16E-01, P=1.93×10−2), lower risks of late-onset Alzheimer’s disease (LOAD) (OR=0.897 (0.810–0.994), P=3.88×10−2), chronic heart failure (CHF) (OR=0.953 (0.913–0.995), P=2. 80×10−2), essential hypertension (OR=0.966 (0.959–0.973), P=1.20×10−22), facial aging (Beta=-1.76E-02, P=2.16×10−9), and early onset chronic obstructive pulmonary disease (COPD) (OR=0.838 (0.755–0.931), P=9.50×10−4) and mildly higher risk of osteoporosis (OR=1.009 (1.006–1.013), P=1.95×10−7). These associations were still significant after excluding SNPs associated with BMI, except for CHF. One potential mechanism of antagonistic pleiotropy is accelerated cell growth and division; therefore, we examined the relationship between menarche and certain cancers. Later age at menarche was associated with lower risks of breast cancer (OR=0.859 (0.764–0.966), P=1.53×10−2) and endometrial cancer (OR=0.896 (0.840–0.956), P=9.02×10−4) compared with early menarche (Fig. 1B).

Genetic associations between exposures and outcomes.

Earlier version: (A), genetic associations between exposures and general aging. Later age at menarche was associated with a lower frailty index and higher parental ages at death. Later age at first birth was associated with lower frailty index and GrimAge and higher parental ages at death. (B), genetic associations between age at menarche and outcomes of organ aging and diseases as well as organ cancer. Later age at menarche was associated with later age at menopause onset and lower risks of LOAD, CHF, essential hypertension, facial aging, early onset COPD, breast cancer, and endometrial cancer. (C), genetic associations between age at first birth and outcomes of organ aging and diseases as well as organ cancer. Later age at first birth was associated with later age at menopause onset and lower risks of type 2 diabetes, CHF, essential hypertension, gastrointestinal or abdominal disease, and cervical cancer. The significant associations were not detected after BMI-related SNPs excluded for outcomes of CHF and cervical cancer. BMI included, BMI-related SNPs included; LOAD, late-onset Alzheimer’s disease; CHF, chronic heart failure; COPD, chronic obstructive pulmonary; GAD, gastrointestinal or abdominal; BMI, body mass index.

Revised version: (A), genetic associations between exposures and aging outcomes. Later age at menarche was associated with a lower frailty index and facial aging, higher parental ages at death, and later menopause. Later age at first birth was associated with lower frailty index and GrimAge, higher parental ages at death, and later menopause. (B), genetic associations between age at menarche and outcomes of organ diseases. Later age at menarche was associated with lower risks of LOAD, CHF, essential hypertension, early onset COPD, breast cancer, and endometrial cancer. (C), genetic associations between age at first birth and outcomes of organ diseases. Later age at first birth was associated with lower risks of type 2 diabetes, CHF, essential hypertension, gastrointestinal or abdominal disease, and cervical cancer. The significant associations were not detected after BMI-related SNPs excluded for outcomes of CHF and cervical cancer. BMI included, BMI-related SNPs included; LOAD, late-onset Alzheimer’s disease; CHF, chronic heart failure; COPD, chronic obstructive pulmonary; GAD,

Later age at first birth was genetically associated with later aging outcomes

Earlier version: Next, we examined the associations between the age of childbirth and age-related outcomes. Compared to early first birth, later age at first birth was associated with lower frailty index (Beta=-6.48E-02, P=5.93×10−11) and GrimAge, a measure of epigenetic aging (Beta=2.67E-01, P=4.37×10−2), and older parental ages at death (Beta=3.43E-02, P=6.88×10−5 for father’s age at death; Beta=3.37E-02, P=1.78×10−4 for mother’s age at death) (Fig. 1A). Similar to results with the age of menarche, later age at first birth was associated with later menopause onset (Beta=2.67E-01, P=3.08×10−4), lower risks of type 2 diabetes (OR=0.890 (0.847– 0.935), P=3.65×10−6), CHF (OR=0.926 (0.866–0.991), P=2.56×10−2), essential hypertension (OR=0.989 (0.984–0.995), P=1.54×10−4), and gastrointestinal or abdominal disease (GAD) (OR=0.991 (0.985–0.996), P=1.29×10−3). Most significant associations remained after the BMI-related SNPs were excluded, except for CHF (Fig. 1C and Table S3). Furthermore, later age at first birth was also significantly associated with a lower risk of cervical cancer (OR=0.999 (0.998–1.000), P=1.52×10−2) but not breast and endometrial cancers (Fig. 1C).

Revised version: Next, we examined the associations between the age of childbirth and age-related outcomes. Compared to early first birth, later age at first birth was associated with lower frailty index (Beta=-6.48E-02, P=5.93×10−11) and GrimAge, a measure of epigenetic aging (Beta=2.67E-01, P=4.37×10−2), and older parental ages at death (Beta=3.43E-02, P=6.88×10−5 for father’s age at death; Beta=3.37E-02, P=1.78×10−4 for mother’s age at death) (Fig. 1A). Similar to results with the age of menarche, later age at first birth was associated with later menopause onset (Beta=2.67E-01, P=3.08×10−4), lower risks of type 2 diabetes (OR=0.890 (0.847– 0.935), P=3.65×10−6), CHF (OR=0.926 (0.866–0.991), P=2.56×10−2), essential hypertension (OR=0.966 (0.957–0.974), P=3.45×10−15), and gastrointestinal or abdominal disease (GAD) (OR=0.975 (0.965–0.985), P=2.02×10−6). Most significant associations remained after the BMI-related SNPs were excluded, except for CHF (Fig. 1A and 1C and Table S3). Furthermore, later age at first birth was also significantly associated with a lower risk of cervical cancer (OR=0.978 (0.975–0.982), P=8.67×10−38) but not breast and endometrial cancers (Fig. 1C).

BMI is an important mediator in significant associations

Earlier version: As BMI is an important modulator of aging 12, we examined its role in explaining these associations. Based on the two-sample MR analyses, significant associations between early menarche and CHF, and between early first birth and CHF and cervical cancer, were not observed after excluding SNPs related to BMI (Fig. 1). To estimate the effect of BMI as the mediator, we further conducted two-step MR. Exposures of age of menarche and age at first birth, and outcomes of frailty index, type 2 diabetes, and CHF, were included in the two-step MR analysis. Later ages of menarche and first birth were associated with lower BMI (Beta=-5.36E-02, P=2.73×10−13; Beta=-4.49E-02, P=6.85×10−5). Higher BMI was associated with higher risks of type 2 diabetes (OR=2.391 (2.243–2.550), P=4.65×10−156), CHF (OR=1.690 (1.570–1.819), P=1.67×10−44), and cervical cancer (OR=1.001 (1.000–1.002), P=3.16×10−2). Significant pleiotropy was detected in the association between BMI and frailty index. Two-step MR showed that BMI significantly mediated the associations between menarche and, type 2 diabetes or CHF (Table S4). BMI also significantly mediated the associations between first birth and, type 2 diabetes, CHF, or cervical cancer (Table S4). As both early menarche and first birth were associated with higher BMI, the BMI had a significant effect to explain the associations partially. Results of colocalization analysis showed 5 SNPs associated with both the exposures and outcomes (Table S5). The genetic correlation results analyzed with LDSC method are listed in Table S6.

Revised version: As BMI is an important modulator of aging 12, we examined its role in explaining these associations. Based on the two-sample MR analyses, significant associations between early menarche and CHF, and between early first birth and CHF and cervical cancer, were not observed after excluding SNPs related to BMI (Fig. 1). To estimate the effect of BMI as the mediator, we further conducted two-step MR. Exposures of age of menarche and age at first birth, and outcomes of frailty index, type 2 diabetes, and CHF, were included in the two-step MR analysis. Later ages of menarche and first birth were associated with lower BMI (Beta=-5.36E-02, P=2.73×10−13; Beta=-4.49E-02, P=6.85×10−5). Higher BMI was associated with higher risks of type 2 diabetes (OR=2.391 (2.243–2.550), P=4.65×10−156), CHF (OR=1.690 (1.570–1.819), P=1.67×10−44), and cervical cancer (OR= 1.019 (1.005–1.033), P=8.18×10−3). Significant pleiotropy was detected in the association between BMI and frailty index. Two-step MR showed that BMI significantly mediated the associations between menarche and, type 2 diabetes or CHF (Table S4). BMI also significantly mediated the associations between first birth and, type 2 diabetes, CHF, or cervical cancer (Table S4). As both early menarche and first birth were associated with higher BMI, the BMI had a significant effect to explain the associations partially. Results of colocalization analysis showed 5 SNPs associated with both the exposures and outcomes (Table S5). The genetic correlation results analyzed with LDSC method are listed in Table S6.

Mechanisms explaining antagonistic pleiotropy

Earlier version: We identified 126 significant SNPs in the genes from the MR analysis for their association between age of menarche or first birth with 15 aging and age-related disease outcomes. These SNPs were then listed according to (i) the number of study outcomes they were associated with and (ii) their frequency of occurrence in the European and global population. We chose to depict only the SNPs that were associated with greater than three outcomes as shown in Fig. 2. The SNP rs2003476 in the CRTC1 gene was found to be associated with seven aging and age-related disease outcomes (frailty index, LOAD, CHF, father’s age at death, mother’s age at death, breast, and endometrial cancer). Previous studies have indicated that CRTC1 transcription domain is linked with extending the lifespan in C. elegans 13 and a decrease in CRTC1 levels was found to be associated with human AD 14. Two other genes involved in the glutathione metabolism pathway, CHAC1 and GGT7, were also found to be associated with four different aging outcomes (CHAC1 was associated with frailty index, diabetes, COPD, and BMI; GGT7 was associated with father’s age at death, hypertension, breast cancer and BMI) (Fig. 2). Importantly, the SNP in CHAC1 is a missense variant, suggesting its impact on protein translation and potential contribution to disease biology. CHAC1 has recently been implicated in age-related macular degeneration 15 and muscle wasting 16. Posttranslational protein-lipid modification processes are known to contribute to aging and age-related diseases 17. N-myristoylation is one such process that is catalyzed by NMT (N-myristoyltransferase) 18, another intriguing gene candidate identified in our analysis. We found an association of NMT1 with three aging outcomes: T2D, CHF, and endometrial cancer. Various myristoylated proteins have been implicated in diverse intracellular signaling pathways 19 including AMP-activated protein kinase (AMPK) 20,21. It has been shown that NMT1 exerts synovial tissue-protective functions by promoting the recruitment of AMPK to lysosomes and inhibiting mTORC1 signaling 20. In addition, drugs targeting NMTs (NMT1 and NMT2) have been suggested to be potent senolytics 22 and a target for treating or preventing various diseases such as cancer 23,24 and heart failure 25. Furthermore, modulating the evolutionarily conserved N-myristoyl transferase NMT1 was shown to extend lifespan in yeast 26. The full table of SNPs and their population frequency and association with aging outcomes is provided in Table S7.

Revised version: We identified 158 significant SNPs in the genes from the MR analysis for their association between age of menarche or first birth with 15 aging and age-related disease outcomes. These SNPs were then listed according to (i) the number of study outcomes they were associated with and (ii) their frequency of occurrence in the European and global population. We chose to depict only the SNPs that were associated with greater than three outcomes as shown in Fig. 2. The SNP rs2003476 in the CRTC1 gene was found to be associated with seven aging and age-related disease outcomes (frailty index, LOAD, CHF, father’s age at death, mother’s age at death, breast, and endometrial cancer). Previous studies have indicated that CRTC1 transcription domain is linked with extending the lifespan in C. elegans 13 and a decrease in CRTC1 levels was found to be associated with human AD 14. Two other genes involved in the glutathione metabolism pathway, CHAC1 and GGT7, were also found to be associated with four different aging outcomes (CHAC1 was associated with frailty index, diabetes, COPD, and BMI; GGT7 was associated with father’s age at death, endometrial cancer and BMI) (Fig. 2).

List and heatmap of SNPs/genes from post-MR analysis showing an association between age at menarche and first birth with three or more aging outcomes.

Earlier version: Genomic region location for each variant is depicted by different characters: intron variant#; 2KB upstream variant$; Missense variant*; Intergenic variant&; Regulatory region variant@. 3’ UTR variant†. Heatmap is representative of an association of each gene/SNP with different aging outcomes. Red bar represents a harmful association (-beta for Father’s DA, Mother’s DA, and menopause; +beta for other outcomes); Blue bar represents a beneficial association (+beta for Father’s DA, Mother’s DA, and menopause; -beta for other outcomes); White bar represents no association. AFB; age at first birth; FI, frailty index; Father’s DA, father’s age at death; Mother’s DA, mother’s age at death; Menopause, age at menopause; LOAD, late-onset Alzheimer’s disease; OS, osteoporosis; T2D, type II diabetes; CHF, chronic heart failure; EH, essential hypertension; FA, facial aging; COPD, chronic obstructive pulmonary disease; GAD, gastrointestinal or abdominal disease; BC, breast cancer; EC, endometrial cancer; BMI, body mass index.

Revised version: Genomic region location for each variant is depicted by different characters: intron variant#; 2KB upstream variant$; Missense variant*; Intergenic variant&; Regulatory region variant@; 3’ UTR variant†; Synonymous Variant∼. Heatmap is representative of an association of each gene/SNP with different aging outcomes. Red bar represents a harmful association (-beta for Father’s DA, Mother’s DA, and menopause; +beta for other outcomes); Blue bar represents a beneficial association (+beta for Father’s DA, Mother’s DA, and menopause; -beta for other outcomes); White bar represents no association. AFB; age at first birth; FI, frailty index; Father’s DA, father’s age at death; Mother’s DA, mother’s age at death; Menopause, age at menopause; LOAD, late-onset Alzheimer’s disease; OS, osteoporosis; T2D, type II diabetes; CHF, chronic heart failure; EH, essential hypertension; FA, facial aging; COPD, chronic obstructive pulmonary disease; GAD, gastrointestinal or abdominal disease; BC, breast cancer; EC, endometrial cancer; CC, cervical cancer; BMI, body mass index.

Added in the revised version: Importantly, the SNP in CHAC1 is a missense variant, suggesting its impact on protein translation and potential contribution to disease biology. CHAC1 has recently been implicated in age-related macular degeneration 15 and muscle wasting 16. Posttranslational protein-lipid modification processes are known to contribute to aging and age-related diseases 17. N-myristoylation is one such process that is catalyzed by NMT (N-myristoyltransferase) 18, another intriguing gene candidate identified in our analysis. We found an association of NMT1 with four aging outcomes: T2D, CHF, essential hypertension, and breast cancer.

Added in the revised version: Various myristoylated proteins have been implicated in diverse intracellular signaling pathways 19 including AMP-activated protein kinase (AMPK) 20,21. It has been shown that NMT1 exerts synovial tissue-protective functions by promoting the recruitment of AMPK to lysosomes and inhibiting mTORC1 signaling 20. In addition, drugs targeting NMTs (NMT1 and NMT2) have been suggested to be potent senolytics 22 and a target for treating or preventing various diseases such as cancer 23,24 and heart failure 25. Furthermore, modulating the evolutionarily conserved N-myristoyl transferase NMT1 was shown to extend lifespan in yeast 26. The full table of SNPs and their population frequency and association with aging outcomes is provided in Table S7.

Canonical signaling pathways analysis

Earlier version: A total of 498 canonical signaling pathways related to 126 SNPs/gene outcomes were identified based on the Ingenuity Pathway Knowledge Base (IPKB). After ranking the identified canonical signaling pathways according to their adjusted P-values, the top 25 canonical signaling pathways with a p-value < 10−2 enriched by SNPs/genes involved in the age of menarche and first birth are represented in Fig. 3A. The top enriched canonical signaling pathways fell into these broader categories: (1) Developmental and cellular signaling pathways such as BMP, IGF-1, and growth hormone signaling; (2) Neuronal signaling and neurological disorders that include amyloid processing, activation of NMDA receptors and postsynaptic events, and glioma signaling; (3) Metabolic signaling pathways such as insulin receptor, mTOR, and leptin signaling; (4) Immunological and inflammatory signaling pathways, which include lymphotoxin β receptor and TREM1 signaling; (5) Cardiovascular and muscular signaling pathways including the role of NFAT in cardiac hypertrophy. In addition to canonical pathways, SNPs/genes were shown to be further associated with diseases and functions. IPA showed a total of 80 diseases and functions associated with age at menarche and first birth. These associated diseases and functions were rated according to their adjusted P-values, and the top 20 enriched categories with a P-value < 10−5 were selected and depicted in Fig. 3B.

Revised version: A total of 563 canonical signaling pathways related to 158 SNPs/gene outcomes were identified based on the Ingenuity Pathway Knowledge Base (IPKB). After ranking the identified canonical signaling pathways according to their adjusted P-values, the top 25 canonical signaling pathways with a p-value < 10−2 enriched by SNPs/genes involved in the age of menarche and first birth are represented in Fig. 3A. The top enriched canonical signaling pathways fell into these broader categories: (1) Developmental and cellular signaling pathways, such as BMP, IGF-1, and growth hormone signaling, and signaling by Activin; (2) Neuronal signaling and neurological disorders, including glioma signaling, amyloid processing, activation of NMDA receptors and postsynaptic events, and neuropathic pain signaling; (3) Metabolic and endocrine signaling pathways, such as mTOR signaling, leptin signaling, PFKFB4 signaling, and melatonin signaling; (4) Immunological and inflammatory pathways, which include lymphotoxin β receptor and IL-22 signaling; (5) Cardiovascular and muscular signaling pathways, including the role of NFAT in cardiac hypertrophy. In addition to canonical pathways, SNPs/genes were shown to be further associated with diseases and functions. IPA showed a total of 78 diseases and functions associated with age at menarche and first birth. These associated diseases and functions were rated according to their adjusted P-values, and the top 20 enriched categories with a P-value < 10−5 were selected and depicted in Fig. 3B.

Earlier version: Ingenuity Pathway Analysis (IPA). SNPs/gene outcomes from MR analysis were subjected to IPA.

Revised version: Ingenuity Pathway Analysis (IPA). SNPs/gene outcomes from MR analysis were subject d to IPA.

Earlier version: (A) Canonical signaling pathways in age at menarche and first birth. The adjusted P-values of the top 25 signaling pathways are listed. (B) Disease and functions in age at first birth and menarche. The adjusted P-value of the top 20 significantly involved diseases and functions are listed. (C) IPA network 1 and (D) IPA network 2. The genes from our post-MR analysis are in bold. Solid lines indicate direct connections, while dotted lines indicate indirect connections (circular arrows mean influence itself). Color coding: pink-receptors, green-enzymes, purple-channels, orange-other proteins.

Revised version: (A) Canonical signaling pathways in age at menarche and first birth. The adjusted P-values of the top 25 signaling pathways are listed. (B) Disease and functions in age at first birth and menarche. The adjusted P-value of the top 20 significantly involved diseases and functions are listed. (C) IPA network 1 and (D) IPA network 2. The genes from our post-MR analysis are in bold. Solid lines indicate direct connections, while dotted lines indicate indirect connections (circular arrows mean influence itself). Color coding: pink-receptors, green-enzymes, blue-transcriptional regulators, red-growth factor, orange-other proteins.

Gene network analysis and drug interactions

Earlier version: Besides pathway and disease associations, the 126 SNPs/genes were further connected to identify 11 functional networks using IPA. These IPA networks were ranked based on their consistency scores. The top 5 networks included a range of 13-17 genes with scores above 25, indicating robust regulatory analysis 27 (Table S8). The top networks were mainly connected to the following functions: Cellular interaction/signaling/development (Fig. 3C), cell death and survival, cancer, gastrointestinal, reproductive (Fig. 3C), neurological (Fig. 3D), and cardiovascular system, development, function, diseases. These results support the notion that dynamic cellular interactions and the development of various vital organ systems influenced by age at menarche and first birth also influence age-related diseases. The two top IPA networks were further outlined in Fig. 3C and 3D. To further explore the drugs targeting genes associated with age at menarche and first birth, we conducted an analysis using ChEMBL28 and DrugBank databases29. This examination revealed that a total of 11 FDA-approved drugs target the identified genes (PRKAG2, SCN2A, DPYD, MC3R, AKT, MAPK, KCNK9, TACR3, HTR4, KCNB2, RXRG) (Table S9).

Revised version: Besides pathway and disease associations, the 158 SNPs/genes were further connected to identify 12 functional networks using IPA. These IPA networks were ranked based on their consistency scores. The top 5 networks included a range of 13-22 genes with scores above 25, indicating robust regulatory analysis 27 (Table S8). The top networks were mainly connected to the following functions: cancer, hematological disease, immunological disease (Fig. 3C), cardiovascular disease, organismal injury and abnormalities, reproductive system development and function (Fig. 3D), and cell-to-cell signaling and interaction. These results support the notion that dynamic cellular interactions and the development of various vital organ systems influenced by age at menarche and first birth also influence age-related diseases. The two top IPA networks were further outlined in Fig. 3C and 3D. To further explore the drugs targeting genes associated with age at menarche and first birth, we conducted an analysis using ChEMBL28 and DrugBank databases29. This examination revealed that a total of 11 FDA-approved drugs target the identified genes (PRKAG2, SCN2A, DPYD, MC3R, AKT, MAPK, KCNK9, TACR3, HTR4, RXRG) (Table S9).

Population validation for associations between age at menarche, or first live birth, or number of births with age-related outcomes

Next, we validated the genetic associations in the population cohort of the UK Biobank. 264,335, 184,481, and 259,758 participants were included for independent variables of age at menarche, age at first live birth, and number of births, respectively (Table S10). Ages at menarche were divided into 5 age groups, <11y, 11-12y, 13-14y, 15-16y, and >16y. The probability of living to 80 years old, having menopause ≥50 years old, risk of diabetes, high blood pressure (HBP), heart failure, COPD, breast cancer, endometrial cancer, and obesity were compared among the 5 age groups of menarche. Each of the outcomes was compared between the highest and lowest education levels and 3 BMI categories (18.5-24.9, 25-29.9, and ≥30) among the 5 age groups (Fig. 4A-I). The ORs for each age group were calculated based on the coefficients (β) in logistic regression after including confounders of smoking, alcohol, education, and BMI. Compared to the group with menarche under 11 years old, which generally had the highest risk, females with menarche at 13-14y had the highest probability of living to 80 years old (OR=1.412, P<0.001) and having menopause ≥50 years old (OR=1.146, P<0.001), and had the lowest risk of diabetes (OR=0.773, P<0.001), heart failure (OR=0.835, P<0.001), and COPD (OR=0.728, P<0.001). The group with menarche at 15-16y had the lowest risks of high blood pressure (HBP) (OR=0.793, P<0.001) and obesity (OR=0.337, P<0.001). Females with menarche >16y had the lowest risks of breast cancer (OR=0.832, P<0.05) and endometrial cancer (OR=0.641, P<0.05) compared to females with menarche <11y (Table S11).

The distributions of outcome age and risk according to age group of menarche and fist live birth as well as number of births before and after correcting confounders.

Condition Baseline, no confounder was corrected. Conditions Highest education and Lowest education, confounders of education, smoking, and drinking were corrected, and bars were painted based on most smoking and drinking situations with the highest and lowest education levels. Conditions Median education with BMI 18.5-24.9, Median education with BMI 25-29.9, and Median education with BMI 30-, confounders of education, smoking, drinking, and BMI were corrected, and bars were painted based on median education and most smoking and drinking situation with 3 BMI categories. BMI, body mass index; HBP, high blood pressure; COPD, chronic obstructive pulmonary disease; AD, Alzheimer’s disease.

Ages at first birth were divided into 5 age groups, <21y, 21-25y, 26-30y, 31-35y, and >35y. The probability of living to 80 years old, having menopause ≥50 years old, risk of diabetes, AD, HBP, heart failure, cerebral infarction, digestive system disease, cervical cancer, and obesity were compared among the 5 age groups of first live birth. Compared to the group with first live birth under 21 years old, which generally had the highest risks, females with first live birth at 21-25y had the highest probability to live to 80 years old (OR=1.667, P<0.001), and females with first live birth at 26-30y had the highest probability to have menopause ≥50 years old (OR=1.184, P<0.001) and lowest risk of cervical cancer (OR=0.587, P<0.01).

Females with first live birth at 31-35y had the lowest risk of diabetes (OR=0.612, P<0.001). Females with first live birth >35y had the lowest risk of having AD (OR=0.474, P<0.01), HBP (OR=0.587, P<0.001), heart failure (OR=0.446, P<0.001), cerebral infarction (OR=0.485, P<0.001), digestive system disease (OR=0.569, P<0.001), and obesity (OR=0.416, P<0.001) (Fig. 4J-S).

Females with the number of births from 0 to 4 were included. Compared to no birth, females with 3 births had the highest probability of living to 80 years old (OR=1.818, P<0.001) and having menopause ≥50 years old (OR=1.562, P<0.001). Females with 4 births had the highest risk of diabetes (OR=1.399, P<0.001), AD (OR=1.676, P<0.001), HBP (OR=1.274, P<0.001), heart failure (OR=1.439, P<0.001), cerebral infarction (OR=1.470, P<0.001), digestive system disease (OR=1.370, P<0.001), and obesity (OR=1.327, P<0.001) but lowest risk of breast cancer (OR=0.869, P<0.001) (Fig. 4T-AC). The significant estimated effect (β) of each group compared to the youngest age groups of menarche and first live birth and 0 birth group and the ability to explain the variability of outcomes in regression analysis are exhibited in Table S11.

Removed in the revised version: To estimate the combined effect of age at menarche and first live birth on diabetes, HBP, heart failure, and BMI≥30, participants were divided into 25 groups according to the menarche and first live birth ages.

Earlier version: Females with menarche<11y and first live birth<21y had higher risks of diabetes (OR=1.881, P<0.001), HBP (OR= 1.436, P<0.05), heart failure (OR= 1.776, P<0.01), and BMI≥30 (OR= 4.417, P<0.001) compared to females with menarche of 13-14y and first live birth of 26-30y (Table S12 and fig. S3).

Revised version: To estimate the combined effect of age at menarche and first live birth on diabetes, HBP, heart failure, and BMI≥30, participants were divided into 25 groups according to the menarche and first live birth ages. Females with menarche<11y and first live birth<21y had higher risks of diabetes (OR=1.881, P<0.001), HBP (OR= 1.436, P<0.05), heart failure (OR= 1.776, P<0.01), and BMI≥30 (OR= 4.417, P<0.001) compared to females with menarche of 13-14y and first live birth of 26-30y (Table S12 and fig. S3).

Discussion

Earlier version: The timing of reproductive milestones such as the age of menarche and age at first birth is partly governed by genetic factors, which may exert lasting impacts on an individual’s health trajectory. Our study finds support for antagonistic pleiotropy by demonstrating that genetic variants that drive early menarche or early pregnancy accelerate several aging outcomes, including parental age, frailty index, and several age-related diseases such as diabetes and dementia. These results were validated in a cohort from the UK biobank, showing that women with early menarche, or early pregnancy, were associated with accelerated aging outcomes and increased risk of several age-related diseases.

Revised version: The timing of reproductive milestones such as the age of menarche and age at first birth is partly governed by genetic factors, which may exert lasting impacts on an individual’s health trajectory. Our study finds support for antagonistic pleiotropy by demonstrating that genetic variants that drive early menarche or early pregnancy accelerate several aging outcomes, including parental age, frailty index, and several age-related diseases such as diabetes and dementia. These results were validated in a cohort from the UK biobank, showing that women with early menarche, or early pregnancy were associated with accelerated aging outcomes and increased risk of several age-related diseases. These results are also consistent with the disposable soma theory that suggests aging as an outcome tradeoff between an organism’s investment in reproduction and somatic maintenance and repair. Both antagonistic pleiotropy theory and the disposable soma theory provide valuable evolutionary frameworks for interpreting how reproductive traits influence long-term health, yet they emphasize distinct biological mechanisms. Antagonistic pleiotropy 2 highlights the temporal trade-off of gene function, proposing that alleles promoting fertility or reproductive success early in life may have detrimental effects in later life. In contrast, the disposable soma theory 30,31 emphasizes a resource allocation trade-off, where investment in reproduction occurs at the cost of somatic maintenance and repair, ultimately accelerating aging and increasing susceptibility to disease. Future research could aim to explore these mechanisms by integrating genetic, metabolic, and longitudinal data, which help understand the interaction between gene-driven effects and resource-driven physiological trade-offs.

Earlier version: Consistent with a recent study suggesting the relation of early first birth with a higher likelihood of frailty 30, our MR analysis provided a strong association between early age of first birth and GrimAge, a DNA methylation based marker to predict epigenetic age. Our findings align with the observed negative genetic correlation between reproductive traits and lifespan that individuals with higher polygenetic scores for reproduction have lower survivorships to age 76 31. Prior studies linking the female reproductive factors with aging are either limited to results obtained from observational studies 32,33 or are limited to few outcomes such as brain disorders 34,35. Besides, emerging evidence suggests that early puberty in males is linked to adverse health outcomes, such as an increased risk of cardiovascular disease, type 2 diabetes, and hypertension in later life 36. A Mendelian randomization study also reported a genetic association between the timing of male puberty and reduced lifespan 37. These findings support the hypothesis that genetic variants associated with delayed reproductive timing in males might similarly confer health benefits or improved longevity, akin to the patterns observed in females. This would suggest that similar mechanisms of antagonistic pleiotropy could operate in males as well.

Revised version: Consistent with a recent study suggesting the relation of early first birth with a higher likelihood of frailty 32, our MR analysis provided a strong association between early age of first birth and GrimAge, a DNA methylation based marker to predict epigenetic age. Our findings align with the observed negative genetic correlation between reproductive traits and lifespan that individuals with higher polygenetic scores for reproduction have lower survivorships to age 76 33. Prior studies linking the female reproductive factors with aging are either limited to results obtained from observational studies 34,35 or are limited to few outcomes such as brain disorders 36,37. Besides, emerging evidence suggests that early puberty in males is linked to adverse health outcomes, such as an increased risk of cardiovascular disease, type 2 diabetes, and hypertension in later life 38. A Mendelian randomization study also reported a genetic association between the timing of male puberty and reduced lifespan 39. These findings support the hypothesis that genetic variants associated with delayed reproductive timing in males might similarly confer health benefits or improved longevity, akin to the patterns observed in females. This would suggest that similar mechanisms of antagonistic pleiotropy could operate in males as well.

Earlier version: Using MR, we identified 126 SNPs associated with early menarche or first birth that significantly influence age-related outcomes. Among the top 25 enriched canonical signaling pathways, those involved in developmental and cellular processes are particularly significant in determining the timing of menarche and first childbirth, with nearly half of these pathways (12 out of 25) falling into this category. These developmental and cellular signaling pathways work together to regulate continuous growth and involution from puberty through menopause 38. Interestingly, known longevity pathways such as insulin-like growth factor 1 (IGF1), growth hormone (GH) signaling, melatonin signaling, and bone morphogenetic protein (BMP) signaling, seem to play a crucial role in regulating the timing of these reproductive events. IGF-1 signaling promotes the growth and development of reproductive organs and tissues 39–41, but is also a conserved modulator of longevity 42–45. Melatonin, known for regulating circadian rhythms, influences the timing of menarche through its effects on the hypothalamic-pituitary-gonadal axis 46 and is postulated to modulate oxidative, inflammatory, and autophagy states 47. Furthermore, different processes in early life such as cell proliferation or differentiation 48, ovarian function, and follicular development 49–51 are dependent on BMP signaling. However, increased BMP signaling also contributes significantly to AD pathology 52 and impairments in neurogenesis and cognitive decline associated with aging 53. Aberrant BMP signaling is also shown to be associated with age-related metabolic and cardiovascular diseases 54,55. Additionally, growth hormone signaling is not only vital for development during childhood and puberty 56 but also impacts longevity 57. Besides signaling pathways, genes in age at menarche and first childbirth were connected in gene-gene interaction networks. In IPA network 1, most genes were connected with either AKT or genes encoding for follicle-stimulating hormone (FSH), luteinizing hormone (LH), leptin receptor (LEPR), and inhibin subunit beta A (INHBA). AKT signaling is implicated in oocyte maturation and embryonic development 58 and mediates pregnancy-induced cardiac adaptive responses 59. AKT signaling also mediates age-related disease pathologies, such as AD 60. Major ovarian functions are controlled by FSH (follicular growth, cellular proliferation, and estrogen production) and LH (oocyte maturation, ovulation, and terminal differentiation of follicles) which in turn are modulated by other ovarian factors such as INHBA (a subunit of activin and inhibin) 61,62. Endocrine alterations at advanced aging or reproductive aging 63 are marked by changes in FSH and LH due to altered feedback resulting from the ovarian decline in sex steroids, inhibin A, and inhibin B production 64–66. In IPA network 2, the hub genes are connected mostly with extracellular signal-regulated kinases (ERK), Similar to AKT, ERK is also a member of the mitogen-activated kinase family (MAPK), a highly conserved signaling pathway that plays a vital role in transducing extracellular signals into a wide range of cellular responses 67,68, maintaining tissue integrity during aging 69, and mediating age-related disease pathology 60,70. These studies underscore the importance of considering that dysregulation or abnormal activation of any of these pathways could contribute to the onset of late-life diseases consistent with antagonistic pleiotropy in humans 2,71.

Revised version: Using MR, we identified 158 SNPs associated with early menarche or first birth that significantly influence age-related outcomes. Among the top 25 enriched canonical signaling pathways, those involved in developmental and cellular processes are particularly significant in determining the timing of menarche and first childbirth, with nearly one third of these pathways (8 out of 25) falling into this category. These developmental and cellular signaling pathways work together to regulate continuous growth and involution from puberty through menopause 40. Interestingly, known longevity pathways such as insulin-like growth factor 1 (IGF1), growth hormone (GH) signaling, melatonin signaling, and bone morphogenetic protein (BMP) signaling, seem to play a crucial role in regulating the timing of these reproductive events. IGF-1 signaling promotes the growth and development of reproductive organs and tissues 41-43, but is also a conserved modulator of longevity 44-47. Melatonin, known for regulating circadian rhythms, influences the timing of menarche through its effects on the hypothalamic-pituitary-gonadal axis 48 and is postulated to modulate oxidative, inflammatory, and autophagy states 49. Furthermore, different processes in early life such as cell proliferation or differentiation 50, ovarian function, and follicular development 51-53 are dependent on BMP signaling. However, increased BMP signaling also contributes significantly to AD pathology 54 and impairments in neurogenesis and cognitive decline associated with aging 55. Aberrant BMP signaling is also shown to be associated with age-related metabolic and cardiovascular diseases 56,57. Additionally, growth hormone signaling is not only vital for development during childhood and puberty 58 but also impacts longevity 59. Besides signaling pathways, genes in age at menarche and first childbirth were connected in gene-gene interaction networks. IPA Network 1 revealed strong connections of our input genes to regulators involved in chromatin remodeling, immune signaling, and proteostasis. Notably, NF-κB, TGF-β, and IL-12—central players in inflammation and immune modulation—emerged as key hubs. While tightly regulated during early life to support immunity and tissue repair, chronic activation of these pathways is associated with inflammaging and immune dysfunction in later life 60,61. Similarly, links to histone H3/H4 and modifiers like USP10 and USP49 suggest that dysregulation of chromatin and ubiquitin-mediated protein turnover may underlie both developmental gene expression programs and age-related epigenetic drift and proteotoxic stress 62,63. The enrichment of E3 ligase-associated nodes (e.g., E3:K48-ubiquitinated substrate, ubiquitin) further supports a central role for proteostasis networks in aging 62,64. These associations reflect the principle of antagonistic pleiotropy—where genes and pathways essential for early-life growth, stress resilience, and immune function become liabilities later in life, contributing to aging and age-related diseases 65. In IPA network 2, most genes were connected with either AKT or genes encoding for follicle-stimulating hormone (FSH), luteinizing hormone (LH), leptin receptor (LEPR), and inhibin subunit beta A (INHBA). AKT signaling is implicated in oocyte maturation and embryonic development 66 and mediates pregnancy-induced cardiac adaptive responses 67. AKT signaling also mediates age-related disease pathologies, such as AD68. Major ovarian functions are controlled by FSH (follicular growth, cellular proliferation, and estrogen production) and LH (oocyte maturation, ovulation, and terminal differentiation of follicles) which in turn are modulated by other ovarian factors such as INHBA (a subunit of activin and inhibin) 69,70. Endocrine alterations at advanced aging or reproductive aging 71 are marked by changes in FSH and LH due to altered feedback resulting from the ovarian decline in sex steroids, inhibin A, and inhibin B production 72-74. These studies underscore the importance of considering that dysregulation or abnormal activation of any of these pathways could contribute to the onset of late-life diseases consistent with antagonistic pleiotropy in humans 2,75.

Earlier version: The thrifty gene hypothesis provides a compelling framework to explain antagonistic pleiotropy, particularly in the context of modern health challenges. The thrifty gene hypothesis suggests that genes favoring efficient energy storage were advantageous in historical periods of food scarcity, helping individuals survive through famines. However, in today’s environment of abundant food, the same genes contribute to obesity and metabolic diseases like diabetes. Similarly, in the context of antagonistic pleiotropy, we hypothesize that genes that enhance early-life reproductive success favor efficient energy storage to support reproductive health. This is supported by our observation that early pregnancy is associated with increasing BMI, a recognized key factor in systemic aging 72,73. In support of this, we found a lack of significant association between early menarche or age at first birth for CHF and cervical cancer after removal of BMI as a confounder in MR analysis. Based on regression analysis at the population level (Fig. 4), higher BMI was associated with higher risks of cardiovascular diseases and diabetes. Thus, an increase in BMI is one of the factors that explains the increased risk of age-related disease seen due to the exposure of early pregnancy or the prevalence of genes that enhance early reproductive success. These results are also consistent with calorie restriction being a robust way to extend healthspan and lifespan in many species 74.

Revised version: The thrifty gene hypothesis provides a compelling framework to explain antagonistic pleiotropy, particularly in the context of modern health challenges. The thrifty gene hypothesis suggests that genes favoring efficient energy storage were advantageous in historical periods of food scarcity, helping individuals survive through famines. However, in today’s environment of abundant food, the same genes contribute to obesity and metabolic diseases like diabetes. Similarly, in the context of antagonistic pleiotropy, we hypothesize that genes that enhance early-life reproductive success favor efficient energy storage to support reproductive health. This is supported by our observation that early pregnancy is associated with increasing BMI, a recognized key factor in systemic aging 76,77. In support of this, we found a lack of significant association between early menarche or age at first birth for CHF and cervical cancer after removal of BMI as a confounder in MR analysis. Based on regression analysis at the population level (Fig. 4), higher BMI was associated with higher risks of cardiovascular diseases and diabetes. Thus, an increase in BMI is one of the factors that explains the increased risk of age-related disease seen due to the exposure of early pregnancy or the prevalence of genes that enhance early reproductive success. These results are also consistent with calorie restriction being a robust way to extend healthspan and lifespan in many species 78.

Earlier version: Several limitations of our study need to be considered. Firstly, all the associations are explored at the genetic level with MR and population level with the UK Biobank. However, to confirm causal relationships, further validation through in vitro and in vivo research is essential. Secondly, though we excluded confounder-related SNPs and addressed potential pleiotropic effects, there may also be potential bias which could be addressed by extending these findings to other populations. Thirdly, significant heterogeneity and pleiotropy remained in some association analyses, which need further exploration. Fourthly, research regarding male-specific reproductive traits and their relationship to aging and health outcomes is needed to compare the difference between male and female aging. Given the increasing age of first childbirth in modern times, the ideal period of first childbirth for both slowing down female aging and benefiting fetal development needs more research 75. In summary, our study underscores the complex relationship between genetic legacies and modern diseases. Understanding these genetic predispositions can inform public health strategies and their relevance to age-related disease outcomes in women of various ethnic groups. For example, interventions could be tailored not only to mitigate the risks associated with early reproductive timing but also to address lifestyle factors exacerbating conditions linked to thrifty genes, like diet and physical activity.

Revised version: Several limitations of our study need to be considered. Firstly, all the associations are explored at the genetic level with MR and population level with the UK Biobank. However, to confirm causal relationships, further validation through in vitro and in vivo research is essential. Secondly, though we excluded confounder-related SNPs and addressed potential pleiotropic effects, there may also be potential bias which could be addressed by extending these findings to other populations. Thirdly, significant heterogeneity and pleiotropy remained in some association analyses, which need further exploration. Fourthly, research regarding male-specific reproductive traits and their relationship to aging and health outcomes is needed to compare the difference between male and female aging. Given the increasing age of first childbirth in modern times, the ideal period of first childbirth for both slowing down female aging and benefiting fetal development needs more research 79. In summary, our study underscores the complex relationship between genetic legacies and modern diseases. Understanding these genetic predispositions can inform public health strategies and their relevance to age-related disease outcomes in women of various ethnic groups. For example, interventions could be tailored not only to mitigate the risks associated with early reproductive timing but also to address lifestyle factors exacerbating conditions linked to thrifty genes, like diet and physical activity.

Methods

All the MR research data is from the public genome-wide association studies (GWAS) databases of the IEU Open GWAS project (https://gwas.mrcieu.ac.uk/) and PubMed. The target analysis is based on public datasets. The population data is from the UK Biobank. No definite personal information was included. No additional ethical approval was required for our study.

Exposure data

2 traits were included as female reproductive activity exposures, involving age at menarche 8 (GWAS ID: ebi-a-GCST90029036) and age at first birth 9 (GWAS ID: ebi-a-GCST90000048).

Outcome data

Earlier version: The aging outcomes involved general aging, organ aging/disease, and organ cancers (Table S1). Frailty index 76 (GWAS ID: ebi-a-GCST90020053), father’s age at death77 (GWAS ID: ebi-a-GCST006700), mother’s age at death 77 (GWAS ID: ebi-a-GCST006699), and DNA methylation GrimAge acceleration 78 (GWAS ID: ebi-a-GCST90014294) were included as general aging outcomes. Specific aging diseases and aging levels included age at menopause onset 8 (GWAS ID: ebi-a-GCST90029037), late-onset Alzheimer’s disease 79 (LOAD) (PMID: 30820047), osteoporosis 80 (GWAS ID: ebi-a-GCST90038656), and type 2 diabetes 81 (GWAS ID: ebi-a-GCST90018926), chronic heart failure 81 (CHF) (GWAS ID: ebi-a-GCST90018806), essential hypertension (GWAS ID: ukb-b-12493), facial aging (GWAS ID: ukb-b-2148), eye aging 82 (PMID: 36975205), cirrhosis (GWAS ID: finn-b-CIRRHOSIS_BROAD), chronic kidney disease 83 (CKD) (GWAS ID: ebi-a-GCST008026), early onset chronic obstructive pulmonary disease (COPD) (GWAS ID: finn-b-COPD_EARLY), and gastrointestinal or abdominal disease (GAD) 80 (GWAS ID: ebi-a-GCST90038597). Organ cancers included breast cancer 81 (GWAS ID: ebi-a-GCST90018799), ovarian cancer 81 (GWAS ID: ebi-a-GCST90018888), endometrial cancer 84 (GWAS ID: ebi-a-GCST006464), and cervical cancer (GWAS ID: ukb-b-8777).

Revised version: The aging outcomes involved general aging, organ aging/disease, and organ cancers (Table S1). Frailty index 80 (GWAS ID: ebi-a-GCST90020053), father’s age at death81 (GWAS ID: ebi-a-GCST006700), mother’s age at death 81 (GWAS ID: ebi-a-GCST006699), and DNA methylation GrimAge acceleration 82 (GWAS ID: ebi-a-GCST90014294) were included as general aging outcomes. Specific aging diseases and aging levels included age at menopause onset 8 (GWAS ID: ebi-a-GCST90029037), late-onset Alzheimer’s disease 83 (LOAD) (PMID: 30820047), osteoporosis 84 (GWAS ID: ebi-a-GCST90038656), and type 2 diabetes 85 (GWAS ID: ebi-a-GCST90018926), chronic heart failure 85 (CHF) (GWAS ID: ebi-a-GCST90018806), essential hypertension (GWAS ID: ukb-b-12493), facial aging (GWAS ID: ukb-b-2148), eye aging 86 (PMID: 36975205), cirrhosis (GWAS ID: finn-b-CIRRHOSIS_BROAD), chronic kidney disease 87 (CKD) (GWAS ID: ebi-a-GCST008026), early onset chronic obstructive pulmonary disease (COPD) (GWAS ID: finn-b-COPD_EARLY), and gastrointestinal or abdominal disease (GAD) 84 (GWAS ID: ebi-a-GCST90038597). Organ cancers included breast cancer 85 (GWAS ID: ebi-a-GCST90018799), ovarian cancer 85 (GWAS ID: ebi-a-GCST90018888), endometrial cancer 88 (GWAS ID: ebi-a-GCST006464), and cervical cancer (GWAS ID: ukb-b-8777).

SNPs selection

Earlier version: We identified single nucleotide polymorphisms (SNPs) associated with exposure datasets with p < 5 × 10−8 85,86. In this case, 249 SNPs and 67 SNPs were selected as eligible instrumental variables (IVs) for exposures of age at menarche and age at first birth, respectively. All selected SNPs for every exposure would be clumped to avoid the linkage disequilibrium (r2 = 0.001 and kb = 10,000). Then we identified whether there were potential confounders of IVs associated with the outcomes based on a database of human genotype-phenotype associations, PhenoScanner V2 87,88 (http://www.phenoscanner.medschl.cam.ac.uk/), with a threshold of p < 1 × 10−5. IVs associated with education, smoking, alcohol, activity, and other confounders related to outcomes would be excluded. During the harmonization process, we aligned the alleles to the human genome reference sequence and removed incompatible SNPs. Subsequent analyses were based on the merged exposure-outcome dataset. We calculated the F statistics to quantify the strength of IVs for each exposure with a threshold of F>10 89. If the effect allele frequency (EAF) was missing in the primary dataset, EAF would be collected from dsSNP (https://www.ncbi.nlm.nih.gov/snp/) based on the population to calculate the F value. As osteoporosis GWAS was processed by BOLT-LMM, we recalculated the beta value for MR analysis. To adjust the beta values (βBOLT) estimated by BOLT-LMM for dichotomous, we applied the following formula to each SNP:

Revised version: We identified single nucleotide polymorphisms (SNPs) associated with exposure datasets with p < 5 × 10−8 89,90. In this case, 249 SNPs and 67 SNPs were selected as eligible instrumental variables (IVs) for exposures of age at menarche and age at first birth, respectively. All selected SNPs for every exposure would be clumped to avoid the linkage disequilibrium (r2 = 0.001 and kb = 10,000). Then we identified whether there were potential confounders of IVs associated with the outcomes based on a database of human genotype-phenotype associations, PhenoScanner V2 91,92 (http://www.phenoscanner.medschl.cam.ac.uk/), with a threshold of p < 1 × 10−5. IVs associated with education, smoking, alcohol, activity, and other confounders related to outcomes would be excluded. During the harmonization process, we aligned the alleles to the human genome reference sequence and removed incompatible SNPs. Subsequent analyses were based on the merged exposure-outcome dataset. We calculated the F statistics to quantify the strength of IVs for each exposure with a threshold of F>10 93. If the effect allele frequency (EAF) was missing in the primary dataset, EAF would be collected from dsSNP (http://www.ncbi.nlm.nih.gov/snp/) based on the population to calculate the F value. As osteoporosis GWAS was processed by BOLT-LMM, we recalculated the beta value for MR analysis. To adjust the beta values (βBOLT) estimated by BOLT-LMM for dichotomous, we applied the following formula to each SNP: The prevalence (P) of the binary outcome in the population is the proportion of individuals with the outcomes. For example, the P value for osteoporosis was calculated based on the GWAS sample size (484,598) and disease cases (7,751). The adjusted beta values for essential hypertension, osteoporosis, GAD, and cervical cancer were calculated in MR analysis.

Removed in the revised version: The prevalence (P) of the binary outcome in the population is the proportion of individuals with the outcomes. The P value for osteoporosis was calculated based on the GWAS sample size (484,598) and disease cases (7,751).

MR analysis

All analysis was performed using R software (version 4.3.1, R Foundation for Statistical Computing, Vienna, Austria) and RStudio software (version 2023.09.0 Posit, PBC, Boston, USA) with TwoSampleMR (version 0.5.7) and MRPRESSO (version 1.0) packages.

Five two-sample MR methods were applied for analysis, including inverse variance weighted model (IVW), MR Egger regression model (MER), weighted median model (WMM), simple mode, and weighted mode. A pleiotropy test was used to check if the IVs influence the outcome through pathways other than the exposure of interest. A heterogeneity test was applied to ensure whether there is a variation in the causal effect estimates across different IVs. Significant heterogeneity test results indicate that some instruments are invalid or that the causal effect varies depending on the IVs used. MRPRESSO was applied to detect and correct potential outliers of IVs with NbDistribution = 10,000 and threshold p = 0.05. Outliers would be excluded for repeated analysis. The causal estimates were given as odds ratios (ORs) and 95% confidence intervals (CI). A leave-one-out analysis was conducted to ensure the robustness of the results by sequentially excluding each IV and confirming the direction and statistical significance of the remaining SNPs.

In two-step MR, we mainly use IVW to access the causal associations between exposure and outcome, exposure and mediator, and mediator and outcome. Body mass index (BMI) (GWAS ID: ukb-b-19953) was applied as a mediator. Similar steps of two-sample MR were repeated in the analysis after excluding confounders. The mediator effect (ME) and direct effect (DE) were calculated. The MEs were calculated with the formula: beta1 × beta2; DEs were calculated according to the formula: beta – (beta1 × beta2), beta stands for the total effect obtained from the primary analysis, beta1 stands for the effect of exposure on the mediator, and beta2 stands for the effect of the mediator on the outcome.

Colocalization analysis

Earlier version: To improve the robustness of our research, colocalization analysis was conducted with packages gwasglue (version 0.0.0.9000), coloc (version 5.2.3), and gassocplot (version 1.0) between the exposures and outcomes revealing significant associations after BMI-related SNPs were excluded. SNPs with significant P values in single SNP OR analysis were set as target SNPs. SNP-level colocalization was conducted with 50kb 90 windows around each target SNP. EAF was set to 0.5 if it was missed in the primary datasets. The Bayesian algorithm in the coloc package generates posterior probabilities (PP) for the hypothesis that both traits are associated and share the same single causal variant at a specific lo91,92cus (H4) 93. SNP with significant results in the colocalization analysis would be removed and two-sample MR analyses would be conducted again.

Revised version: To improve the robustness of our research, colocalization analysis was conducted with packages gwasglue (version 0.0.0.9000), coloc (version 5.2.3), and gassocplot (version 1.0) between the exposures and outcomes revealing significant associations after BMI-related SNPs were excluded. SNPs with significant P values in single SNP OR analysis were set as target SNPs. SNP-level colocalization was conducted with 50kb 94 windows around each target SNP. EAF was set to 0.5 if it was missed in the primary datasets. The Bayesian algorithm in the coloc package generates posterior probabilities (PP) for the hypothesis that both traits are associated and share the same single causal variant at a specific lo95,96cus (H4) 97. SNP with significant results in the colocalization analysis would be removed and two-sample MR analyses would be conducted again.

Genetic correlation analysis

Earlier version: Genetic correlation analysis helps clarify whether there is a shared genetic architecture between two traits, which can support the interpretation of our MR results. LDSC software, a robust framework to estimate genetic correlation using GWAS summary statistics, was applied to calculate the genetic correlation 91,92. GWAS summary statistics were filtered with HapMap3 ref. Variants, which were repeated or had a minor allele frequency (MAF) ≤0.01 were excluded. Genetic Correlation (rg) was calculated for all pairwise comparisons among age at menarche, age at first birth, BMI, and age at menopause onset using LDSC method.Target Analysis

Revised version: Genetic correlation analysis helps clarify whether there is a shared genetic architecture between two traits, which can support the interpretation of our MR results. LDSC software, a robust framework to estimate genetic correlation using GWAS summary statistics, was applied to calculate the genetic correlation 95,96. GWAS summary statistics were filtered with HapMap3 ref. Variants, which were repeated or had a minor allele frequency (MAF) ≤0.01 were excluded. Genetic Correlation (rg) was calculated for all pairwise comparisons among age at menarche, age at first birth, BMI, and age at menopause onset using LDSC method.

Added in the revised version: Target Analysis

Earlier version: We employed the Ingenuity Pathway Analysis (IPA) software (version 01-22-01; Ingenuity Systems; QIAGEN) to investigate various gene-related aspects associated with age at first birth and menarche. IPA is a widely used bioinformatics tool for interpreting high-throughput data 94. Briefly, the SNPs/genes from MR analysis were uploaded into Qiagen’s IPA system for core analysis and then the outcome was overlaid with the global molecular network in the Ingenuity pathway knowledge base (IPKB). IPA was performed to identify canonical pathways, diseases, and functions, and to investigate gene networks. Additionally, the Chemical Biology Database (ChEMBL) (https://www.ebi.ac.uk/chembl/), a large-scale database of bioactive molecules for drug discovery 28, and the DrugBank database 29, a comprehensive database on FDA-approved drugs, drug targets, mechanisms of action, and interactions (https://go.drugbank.com/), were used to identify candidate genes targeted by approved drugs in clinical trials or under current development phases.

Revised version: We employed the Ingenuity Pathway Analysis (IPA) software (version 01-22-01; Ingenuity Systems; QIAGEN) to investigate various gene-related aspects associated with age at first birth and menarche. IPA is a widely used bioinformatics tool for interpreting high-throughput data 98. Briefly, the SNPs/genes from MR analysis were uploaded into Qiagen’s IPA system for core analysis and then the outcome was overlaid with the global molecular network in the Ingenuity pathway knowledge base (IPKB). IPA was performed to identify canonical pathways, diseases, and functions, and to investigate gene networks. Additionally, the Chemical Biology Database (ChEMBL) (http://www.ebi.ac.uk/chembl/), a large-scale database of bioactive molecules for drug discovery 28, and the DrugBank database 29, a comprehensive database on FDA-approved drugs, drug targets, mechanisms of action, and interactions (https://go.drugbank.com/), were used to identify candidate genes targeted by approved drugs in clinical trials or under current development phases.

Population validation

Earlier version: The MR results were further validated based on the UK Biobank with package stats (version 4.3.1). In addition to age at menarche and first live birth, the number of births is listed as an independent variable for the validation analysis. Based on the significant associations, regression analysis was adopted to explore the effect of independent variables and related confounders on outcomes. Ages at menarche were divided into 5 age groups, <11y, 11-12y, 13-14y, 15-16y, and >16y. Ages at first live birth were divided into 5 age groups, <21y, 21-25y, 26-30y, 31-35y, and >35y. Females with the number of births from 0 to 4 were included. Education was divided into 7 levels based on the degrees. Smoking was divided into 2 categories, ever smoking or not.

Revised version: The MR results were further validated based on the UK Biobank with package stats (version 4.3.1). In addition to age at menarche and first live birth, the number of births is listed as an independent variable for the validation analysis. Based on the significant associations, regression analysis was adopted to explore the effect of independent variables and related confounders on outcomes. Ages at menarche were divided into 5 age groups, <11y, 11-12y, 13-14y, 15-16y, and >16y. Ages at first live birth were divided into 5 age groups, <21y, 21-25y, 26-30y, 31-35y, and >35y. Females with the number of births from 0 to 4 were included. Education was divided into 7 levels based on the degrees. Smoking was divided into 2 categories, ever smoking or not. Drinking was divided into 3 categories, never, previous, and current drinking. BMI was divided into four categories, <18.5, 18.5-24.9, 25-29.9, and ≥30 based on the average value of 4 records.

Earlier version: Drinking was divided into 3 categories, never, previous, and current drinking. BMI was divided into four categories, <18.5, 18.5-24.9, 25-29.9, and ≥30 based on the average value of 4 records. The outcome results were preprocessed based on the baseline information collection, follow-up outcome, surgical history, and ICD-10 diagnosis summary. Logistic regression was used for the analysis of disease risks (categorical variables). The coefficients (β) in logistic regression were log-odds ratios. Although the continuous variables exhibited mild to moderate skewed distributions, considering the large sample size, linear regression was used for the analysis of age at death, menopause age, and BMI in the first step. Then logistic regression was applied to confirm the robustness of the results with the outcome of death age ≥ 80, menopause age ≥ 50, and BMI ≥ 30. Confounders for regression analysis I include education, smoking, and drinking. Confounders for regression analysis II include education, smoking, drinking, and BMI. Variance Inflation Factors (VIF) were calculated for each predictor in regression analysis to avoid multicollinearity. Values of VIF were less than 1.5 for each predictor. For outcomes of diabetes, HBP, heart failure, and BMI≥30, the combined effect of age at menarche and first birth was analyzed. 179,821 participants were divided into 25 groups according to the menarche and first live birth age groups. All other groups were compared to the group with menarche of “13-14y” and the first birth of “26-30y” in logistic regression (confounders including education, smoking, drinking, and BMI for diabetes, HBP, and heart failure; including education, smoking, and drinking for BMI≥30).

Revised version: The outcome results were preprocessed based on the baseline information collection, follow-up outcome, surgical history, and ICD-10 diagnosis summary. Logistic regression was used for the analysis of disease risks (categorical variables). The coefficients (β) in logistic regression were log-odds ratios. Although the continuous variables exhibited mild to moderate skewed distributions, considering the large sample size, linear regression was used for the analysis of age at death, menopause age, and BMI in the first step. Then logistic regression was applied to confirm the robustness of the results with the outcome of death age ≥ 80, menopause age ≥ 50, and BMI ≥ 30. Confounders for regression analysis I include education, smoking, and drinking. Confounders for regression analysis II include education, smoking, drinking, and BMI. Variance Inflation Factors (VIF) were calculated for each predictor in regression analysis to avoid multicollinearity. Values of VIF were less than 1.5 for each predictor. For outcomes of diabetes, HBP, heart failure, and BMI≥ 30, the combined effect of age at menarche and first birth was analyzed. 179,821 participants were divided into 25 groups according to the menarche and first live birth age groups. All other groups were compared to the group with menarche of “13-14y” and the first birth of “26-30y” in logistic regression (confounders including education, smoking, drinking, and BMI for diabetes, HBP, and heart failure; including education, smoking, and drinking for BMI≥30).